# Aula 1 - Exercícios Extra - Sintaxe de Python e NumPy Array
### Consulta a internet (Google, Stackoverflow e etc) liberada!!

#### Fonte dos exercícios:
1 - https://www.w3resource.com/python-exercises/python-basic-exercises.php

2 - https://www.practicepython.org/    

### 1.11 Escreva uma função em Python receba como parâmetro duas listas de numeros inteiros e retorne os elementos que são comuns entre elas (sem repetição). Use a sintaxe inline (list comprehension).
##### Exemplo: a = [1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89] e b = [1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13]
##### Resultado esperado: 

<!-- grade_id: cell-a2f686fdd79e4f51, points: 0, grade: false, solution: true -->
```python
# Dica: aninhe na mesma instrução o loope (for) e a condição (if)
def intersecao(a, b):
    """retorne os valor presentes nas duas listas
    """
    ### BEGIN SOLUTION
    return [i for i in a if i in b]
    ### END SOLUTION
```

```python
## Faça seus testes aqui

```

<!-- grade_id: cell-5412543179b17057, points: 1, grade: true, solution: false -->
```python
# Validação para nota
a = [1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89] 
b = [1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13]
assert intersecao(a,b) == [1, 1, 2, 3, 5, 8, 13]
```

### 1.12 Escreva uma função que receba uma lista de inteiros como parâmetro e retorne um dicionário com as chaves sendo os elementos da lista, e os valores sendo os elementos da mesma lista, mas na ordem inversa.
#### Exemplo: para a lista [1, 2, 3, 4, 5]
#### Resultado esperado: {1: 5, 2: 4, 3: 3, 4: 2, 5: 1}

<!-- grade_id: cell-2e044a5f0541da37, points: 0, grade: false, solution: true -->
```python
# Dica: use a função zip. Ela retorna itens agrupados dois a dois.
def create_dict(lista):
    """retorne o maior item da lista
    """
    ### BEGIN SOLUTION
    mapping = dict(zip(lista, reversed(lista)))
    return mapping
    ### END SOLUTION
```

```python
## Faça seus testes aqui


```

<!-- grade_id: cell-5bd4079b37b6a7f3, points: 1, grade: true, solution: false -->
```python
# Validação para nota
assert create_dict([1, 2, 3, 4, 5]) == {1: 5, 2: 4, 3: 3, 4: 2, 5: 1}
assert create_dict([-100,-1580,-1,-9100,-77]) == {-100: -77, -1580: -9100, -1: -1, -9100: -1580, -77: -100}
```

### 1.13 Escreva uma função que receba como entrada uma string conforme apresentada a seguir e retorne uma lista contendo a população dos 3 Estados brasileiros mais populosos.

"""Acre=881935; Alagoas=3337357; Amapá=845731; Amazonas=4144597; Ceará=9132078;
 Distrito Federal=3015268; Espírito Santo=4018650; Goiás=7018354; Maranhão=7075181;
 Mato Grosso=3484466; Mato Grosso do Sul=2778986; Minas Gerais=21168791;
 Paraíba=4018127; Paraná=11433957; Pernambuco=9557071; Piauí=3273227; Rio de Janeiro=17264943;
 Rio Grande do Norte=3506853; Rio Grande do Sul=11377239; Rondônia=1777225;
 Roraima=605761; Santa Catarina=7164788; São Paulo=45919049; Sergipe=2298696;
 Tocantins=1572866; Bahia=14873064; Pará=8602865"""

#### Nesta string, a informação de cada Estado está separada por um ;
 

<!-- grade_id: cell-5420c3983940f212, points: 0, grade: false, solution: true -->
```python
# Dica 1: A função split quebra uma string numa lista baseado no caracter 
#         de separação especificado. 
#         Exemplo: '1,2,3'.split(',') retorna a lista ['1','2','3']
# Dica 2: a função .sort de uma lista ordena a referida lista


def top3_estados(lista_texto):
    """retorne os 3 estados de maior população. O parâmetro de entrada deve ser uma string
         conforme apresentado na descrição do exercício.
    """
    ### BEGIN SOLUTION
    estados = {}
    for l in lista_texto.split(";"):
        ep = l.split("=")
        estados[ep[0]] = int(ep[1])
    pop = list(estados.values())
    pop.sort(reverse=True)
    return pop[0:3]
    ### END SOLUTION
```

```python
## Faça seus testes aqui


```

<!-- grade_id: cell-f86ec02e2bcb84c6, points: 1, grade: true, solution: false -->
```python
# Validação para nota
lista_estados="""Acre=881935; Alagoas=3337357; Amapá=845731; Amazonas=4144597; Ceará=9132078;\
 Distrito Federal=3015268; Espírito Santo=4018650; Goiás=7018354; Maranhão=7075181;\
 Mato Grosso=3484466; Mato Grosso do Sul=2778986; Minas Gerais=21168791;\
 Paraíba=4018127; Paraná=11433957; Pernambuco=9557071; Piauí=3273227; Rio de Janeiro=17264943;\
 Rio Grande do Norte=3506853; Rio Grande do Sul=11377239; Rondônia=1777225;\
 Roraima=605761; Santa Catarina=7164788; São Paulo=45919049; Sergipe=2298696;\
 Tocantins=1572866; Bahia=14873064; Pará=8602865"""

assert top3_estados(lista_estados) == [45919049, 21168791, 17264943]
```

### 1.14 Escreva uma função Python que aceite como entrada uma string de palavras separadas por vírgula e retorne uma string com as palavras em ordem alfabética e separadas por vírgula. 
#### Exemplo:  Para t="hoje,nota,teste,bolsa"
#### Resultado Esperado: "bolsa,hoje,nota,teste"

<!-- grade_id: cell-bb9537795a06bd7e, points: 0, grade: false, solution: true -->
```python
# Dica: use as funções .split e .join presente nas strings
def sort_words(words):
    """retorne uma string ordenada de palavras separadas por vírgula
    """
    ### BEGIN SOLUTION
    lw = words.split(',')
    lw.sort()
    return ",".join(lw)
    ### END SOLUTION
```

```python
## Faça seus testes aqui


```

<!-- grade_id: cell-61d03fd9d1b0eb4e, points: 1, grade: true, solution: false -->
```python
# Validação para nota
assert sort_words("hoje,nota,teste,bolsa") == "bolsa,hoje,nota,teste"

```